## **Write your answers directly in the cells on this notebook.**

### Show the entire process, from the initial fractions to the final result. 
The probability results must be a decimal value between 0.0 and 1.0. Round your number to the third decimal (e.g 0.574).**

Consider the following boxes (two boxes (red/blue) with two types of balls (green/orange)):
![image.png](attachment:image.png)


#### Exercise 1: which is the probability of selecting the red box?

<!-- grade_id: cell-7969bda5a10dc7ac, points: 0, grade: false, solution: true -->
```python
sol = 0
import random
import numpy as np
float(0.3453453453)
# YOUR CODE HERE
redbox = ["G","G","O","O","O","O","O","O"]
bluebox = ["O","G","G","G"]
box = (redbox,bluebox)
#k = random.choice(box)
sol = 1/len(box)
sol = round(sol,3)
sol
```

Output:
```
0.5
```

```python
'''totalA = 12 # total number of A transistors
totalB = 18 # total number of B transistors
choicelist = list("A" * totalA + "B" * totalB) # list containing transitors to correct proportion

def transistor():
    return random.choice(choicelist) # pick a random transistor from list
transistor()'''
```

Output:
```
'A'
```

<!-- grade_id: cell-2b9f005060121b41, points: 1, grade: true, solution: false -->
```python
assert sol == 0.5
```

#### Exercise 2: if you want to run out of balls both boxes at the same time, how often should you select the blue box?

<!-- grade_id: cell-10a736364bdf6ef5, points: 0, grade: false, solution: true -->
```python
sol = 0
# YOUR CODE HERE
sol = len(bluebox)/(len(redbox)+len(bluebox))

```

<!-- grade_id: cell-d06dc523a3f300ce, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.3333333333333333) < 0.01
```



```
# Bu, kod olarak biçimlendirilmiştir
```

#### Exercise 3: which is the probability of picking up a green ball if you are using the red box?

<!-- grade_id: cell-335a720d5716ad4b, points: 0, grade: false, solution: true -->
```python
sol = 0
# YOUR CODE HERE
sol = redbox.count("G")/len(redbox)
sol
```

Output:
```
0.25
```

```python
random.choices(redbox, k=4)
```

Output:
```
['O', 'O', 'O', 'G']
```

<!-- grade_id: cell-3cc76e304b24df43, points: 1, grade: true, solution: false -->
```python
assert sol == 0.25
```

#### Exercise 4: which is the probability of picking up an orange ball (any of the boxes)?

<!-- grade_id: cell-6be260afb8041372, points: 0, grade: false, solution: true -->
```python
sol = 0
# YOUR CODE HERE
sol = (box[:][1].count("O")+box[:][0].count("O"))/(len(redbox)+len(bluebox))
sol
```

Output:
```
0.5833333333333334
```

<!-- grade_id: cell-57ee289d4e5df3f0, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.5833333333333334) < 0.01
```

#### Optional Exercise 5: given that you picked up a green ball, which is the probability of having selected the red box?

<!-- grade_id: cell-9c4af4d3ec207bb2, points: 0, grade: false, solution: true -->
```python
# No need to use the Bayes Theorem as it has not been explained in class
sol = 0
# YOUR CODE HERE
sol = (1/len(box))*(redbox.count("G")/len(redbox))
sol 
```

Output:
```
0.125
```

<!-- grade_id: cell-10164afff838da6f, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.125) < 0.3
```

#### Optional Exercise 6: I have chosen a ball from a box and I picked up a green one. Which is the probability of having used the blue box?

```python
sol = (1/len(box))*(bluebox.count("G")/len(bluebox))
```

<!-- grade_id: cell-ac2e94b60b1ed6fc, points: 0, grade: false, solution: true -->
```python
# No need to use the Bayes Theorem as it has not been explained in class
sol = 0
# YOUR CODE HERE
sol = (1/len(box))*(bluebox.count("O")/len(bluebox))


```

<!-- grade_id: cell-1ca0be55ea9d081c, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.375) < 0.9
```

Now consider a standard deck of cards (the one used in Poker):
![image.png](attachment:image.png)

#### Exercise 7: which is the probability of drawing an Ace?

<!-- grade_id: cell-f7a947cdf92e17c3, points: 0, grade: false, solution: true -->
```python
# Create function that returns probability percent rounded to three decimal places

def event_probability(event_outcomes, sample_space):
    probability = 0
    event_outcomes = 4
    sample_space = 52

```

<!-- grade_id: cell-adb42109bd9b7996, points: 1, grade: true, solution: false -->
```python
# Using the previously defined function: Determine the probability of drawing an Ace
assert abs(sol-0.077 ) < 0.01
```

```python
a = [1,2,3,4,5,6,7,8,9,10,11,12,13]
s = ["heart","diamond","spades","club"]
x = {"heart":a,s[1]:a,s[2]:a,s[3]:a}
x
```

Output:
```
{'club': [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13],
 'diamond': [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13],
 'heart': [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13],
 'spades': [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]}
```

```python
listt = ["a","j","k","q",2,3,4,5,6,7,8,9]
k = "H".join(str(e) for e in list)
```

Output:
```
'aHjHkHqH2H3H4H5H6H7H8H9'
```

#### Exercise 8: which is the probability of drawing a card that is a Heart?

<!-- grade_id: cell-e261cc3812382a18, points: 0, grade: false, solution: true -->
```python
# Using the previously defined function: Determine the probability of drawing a heart

# YOUR CODE HERE
n_o_hearts = 13
n_o_cards = 52
sol = n_o_hearts/n_o_cards

```

<!-- grade_id: cell-2bcb2371e498a429, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.25 ) < 0.01
```

```python
import itertools, random
# make a deck of cards
decko = list()
deck = list(itertools.product(['A', '2', '3', '4', '5', '6', '7', '8', '9', '10', 'J', 'Q', 'K'],
                              ['Spade','Heart','Diamond','Club']))
deck

```

```python
N = 52
#sol = deck[0].count("Heart")/len(deck)

ele = []
random.shuffle(deck)
for i in range(N):
  aces = [d[0] for d in deck[0:1]].count('Heart')
  ele = ele.append(i)
ele


```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-243-032cdd4f7de2> in <module>()
      6 for i in range(N):
      7   aces = [d[0] for d in deck[0:1]].count('Heart')
----> 8   ele = ele.append(i)
      9 ele
     10 

AttributeError: 'NoneType' object has no attribute 'append'
```

```python
N = 100000
double_aces = 0
for hands in range(N):
    # shuffle the cards
    random.shuffle(deck)
    aces = [d[0] for d in deck[0:2]].count('A')
    if aces == 1:
        double_aces+=1
prob = double_aces/N
prob
```

Output:
```
0.117666
```

```python
random.shuffle(deck)
```

```python
N = 1000
double_aces = 0
for hands in range(N):
    # shuffle the cards
    random.shuffle(deck)
    aces = [d[0] for d in deck[0:2]].count('Diamond')
    if aces == 1:
        print("HEArt")
    print("HEArt")
```

#### Exercise 9: which is the probability of drawing a face card (such as Jacks, Queens, or Kings)?

<!-- grade_id: cell-fc4fcd733d3e1ca7, points: 0, grade: false, solution: true -->
```python
# Using the previously defined function: Determine the probability of drawing a face card

# YOUR CODE HERE
sol= 12/52
```

<!-- grade_id: cell-f17c3aa2f33022be, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.231 ) < 0.01
```

#### Exercise 10: which is the probability of drawing a face card which is also a Heart?

<!-- grade_id: cell-c50a860ddb2f18f9, points: 0, grade: false, solution: true -->
```python
# Using the previously defined function: Determine the probability of drawing the jack, queen or king  of hearts
# YOUR CODE HERE
sol = 3/52
```

<!-- grade_id: cell-dcd84807c13f2154, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.058 ) < 0.01
```

#### Exercise 11: which is the probability of drawing an Ace on the second draw, if the first card drawn was a King?

<!-- grade_id: cell-c953e962deacd22d, points: 0, grade: false, solution: true -->
```python
# Using the previously defined function: Determine the probability of drawing an Ace after drawing a King on the first draw


# YOUR CODE HERE
sol = 4/51
```

<!-- grade_id: cell-c0fd1832c0fa3b01, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.078 ) < 0.01
```

#### Exercise 12: which is the probability of drawing an Ace on the second draw, if the first card drawn was an Ace?

<!-- grade_id: cell-899b031f866ba3b2, points: 0, grade: false, solution: true -->
```python
# Using the previously defined function: Determine the probability of drawing an Ace after drawing an Ace on the first draw

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9109c59db8e329c8, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.059) < 0.01
```

Now consider the following situation:

You are playing Poker in the [Texas Holdem variant](https://en.wikipedia.org/wiki/Texas_hold_%27em). In case you are not familiar, it is a variant of poker in which each player has two cards and there is a set of community cards to play from. 
![image.png](attachment:image.png)

#### Exercise 13: what is the probability that the next card drawn will be a Diamond card?

<!-- grade_id: cell-17b8aec4013e6b9c, points: 0, grade: false, solution: true -->
```python
# Sample Space
cards = 52
player_cards = 2
turn_community_cards = 4
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ebe4abfecd5e8690, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.196) < 0.01
```

Now consider the following situation:
![image.png](attachment:image.png)

#### Exercise 14: which is the probability that with the next card drawn there will be five cards in sequential order? 

(Notice that any Eight ( 8, 9, 10, Jack, Queen) or any King (9, 10, Jack, Queen, King) will complete the straight)

<!-- grade_id: cell-f05c41e0e869e186, points: 0, grade: false, solution: true -->
```python
# Sample Space
cards = 52
player_cards = 2
turn_community_cards = 4

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e46db60e9b2dfa4f, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.174) < 0.01
```

## Now, we keep playing Texas Holdem 
![](https://media.giphy.com/media/EQBsWLfdxHW01Y1OHK/giphy.gif)
### But looking only at the cards we could have in our hand
### (2 possible drawing events, ignore community cards)
#### Exercise 15: which is the probability of drawing a heart OR drawing a club?

<!-- grade_id: cell-5b6c82ce08d8c47f, points: 0, grade: false, solution: true -->
```python
# Calculate the probability of drawing a heart or a club
# YOUR CODE HERE
raise NotImplementedError()
sol

```

<!-- grade_id: cell-dedacb0828533e9c, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.755) < 0.01
```

#### Exercise 16: which is the probability of drawing an ace, a king or a queen?

<!-- grade_id: cell-0b50174237f37b03, points: 0, grade: false, solution: true -->
```python
# Calculate the probability of drawing an ace, king, or a queen
# YOUR CODE HERE
sol = 12/52
```

<!-- grade_id: cell-1c3bec1e6a4f0726, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.41171499999999994) < 0.01
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-360-f6f74afa83b0> in <module>()
----> 1 assert abs(sol-0.41171499999999994) < 0.01

AssertionError:
```

#### Exercise 17: which is the probability of drawing a heart or an ace?

<!-- grade_id: cell-3ff5cdeedc739059, points: 0, grade: false, solution: true -->
```python
# Calculate the probability of drawing a heart or an ace
# YOUR CODE HERE
sol = (13+4-1)/52
sol
```

Output:
```
0.3076923076923077
```

<!-- grade_id: cell-a74509b22cdaf4bd, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.525288) < 0.01
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-358-de876bc4c874> in <module>()
----> 1 assert abs(sol-0.525288) < 0.01

AssertionError:
```

#### Exercise 18: which is the probability of drawing a red card or drawing a face card?

<!-- grade_id: cell-9c43e1f91aac785f, points: 0, grade: false, solution: true -->
```python
# Calculate the probability of drawing a red card or a face card

# YOUR CODE HERE
sol = (26+12-6)/52
sol
```

Output:
```
0.6153846153846154
```

<!-- grade_id: cell-47f891a62515e8d4, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.856395) < 0.01
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-356-a5cb34159255> in <module>()
----> 1 assert abs(sol-0.856395) < 0.01

AssertionError:
```

#### Exercise 19: which is the probability of drawing an Ace from a deck of cards, replacing it, reshuffling the deck, and drawing another Ace?

<!-- grade_id: cell-18d5038a32ffeac9, points: 0, grade: false, solution: true -->
```python
# Sample Space
cards = 52

# Outcomes
aces = 4

# YOUR CODE HERE
sol = (aces/cards)*(aces/cards)
```

<!-- grade_id: cell-536e3de73c88e083, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.005929) < 0.001
```

#### Exercise 20: which is the probability of being dealt two Aces (drawing one Ace after the other starting with a full deck)?

<!-- grade_id: cell-9ec7e543991e2e7c, points: 0, grade: false, solution: true -->
```python
# Sample Space first draw
cards = 52

# Outcomes first draw
aces = 4

# YOUR CODE HERE
sol = round((aces/cards)*((aces-1)/(cards-1)),3)
```

<!-- grade_id: cell-0651acc2d4c97ad3, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.004542999999999999) < 0.001
```

### Permutations and combinations

#### Exercise 21: How many different 5-letter arrangements are there of the letters in the word morse?

<!-- grade_id: cell-5d16f7fa95eb0955, points: 0, grade: false, solution: true -->
```python
import math
sol = 0
# YOUR CODE HERE
sol = 5*4*3*2*1
```

<!-- grade_id: cell-290944fdf9fb212b, points: 1, grade: true, solution: false -->
```python
assert sol == 120
```

#### Exercise 22: How many different seven-letter arrangements of the letters in the word HEXAGON can be made if each letter is used only once?

<!-- grade_id: cell-2b2af5a344e47a8b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE should go with 7!
sol = 7*6*5*4*3*2*1
```

<!-- grade_id: cell-9eed40b79a9bc42b, points: 1, grade: true, solution: false -->
```python
assert sol == 5040
```

#### Exercise 23: Two cards are drawn at random from a standard deck of cards, without replacement. Find the probability of drawing an 8 and a queen in that order.

<!-- grade_id: cell-0772318d22f07226, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
sol = (4/52)*(4/51)
```

<!-- grade_id: cell-a85db4f2cf65b50d, points: 1, grade: true, solution: false -->
```python
assert abs(sol-0.006006) < 0.001
```

```python

```